// Copyright (C) 2025 Cruftless project. Test262-shaped fixture.
// Mirrors test262/test/built-ins/Array/prototype/map/length.js.
/*---
esid: sec-array.prototype.map
description: Array.prototype.map.length is 1
includes: []
---*/
assert.sameValue(Array.prototype.map.length, 1, "Array.prototype.map.length is 1");
